<h1>Problem 1</h1>
<p>
    <b>This problem tests your basic understanding of loops, lists, and appending to a list</b>
</p>
A python list contains daily sales (in dollars) at a store over the last n days. Write a python function that returns a list of tuples where each tuple is the day number and the sales value for all instances where the sales were greater than 100.00

Test your program with the following example:

<pre>
sales = [102.3,98.3,135.64,107.12,88.43,198.16]
</pre>
Your program should return:

<pre>
[(0, 102.3), (2, 135.64), (3, 107.12), (5, 198.16)]
</pre>


<!-- grade_id: prob1_solution, points: 0, grade: false, solution: true -->
```python
def high_sales(sales):
    # YOUR CODE HERE
    result = []
    for i, sal in enumerate(sales):
        if sal > 100.00:
            result.append((i,sal))
    return result


    # raise NotImplementedError()
```

<!-- grade_id: prob1_public_test1, points: 0.25, grade: true, solution: false -->
```python
# Test case 1
sales = [102.3,98.3,135.64,107.12,88.43,198.16]
assert high_sales(sales) == [(0, 102.3), (2, 135.64), (3, 107.12), (5, 198.16)], high_sales(sales)
```

<!-- grade_id: prob1_test1, points: 0.75, grade: true, solution: false -->
```python

```

<h1>Problem 2</h1>

ParapetShop keeps track of its inventory in a dictionary keyed by product. Products can be in multiple locations so the value associated with each key is another dictionary with store name’s as
the key and product quantity as the value. 

For example:

<pre>
inventory = {‘widgets’:{‘paramus’:200,’yonkers’:400,’chappaqua’:323},
{‘spam’:{‘paramus’:450,’yonkers’:800,’chappaqua’:1287}}
</pre>
Write a function that takes an inventory list and an item name as input and returns the total quantity of that item across all stores. If the item is not in the store, the function should return 0. For example:

<pre>
total_quantity(inventory,"widgets")
</pre>
should return:

<pre>
923
</pre>
and
<pre>
total_quantity(inventory,"eggs")
</pre>
should return:

<pre>
0
</pre>

<!-- grade_id: prob2_solution, points: 0, grade: false, solution: true -->
```python
def total_quantity(inventory,item):
    # YOUR CODE HERE
    for key, val in inventory.items():
        if key == item:
            return sum(val.values())
    return 0
    # raise NotImplementedError()
```

<!-- grade_id: prob2_public_test1, points: 0.25, grade: true, solution: false -->
```python
# Test case 1
inventory = {'widgets':{'paramus':200,'yonkers':400,'chappaqua':323},
             'spam':{'paramus':450,'yonkers':800,'chappaqua':1287}}

assert total_quantity(inventory,"widgets") == 923, total_quantity(inventory,"widgets")
```

<!-- grade_id: prob2_public_test2, points: 0.25, grade: true, solution: false -->
```python
# Test case 2
assert total_quantity(inventory,"eggs") == 0, total_quantity(inventory,"eggs")
```

<!-- grade_id: prob2_test1, points: 0.5, grade: true, solution: false -->
```python

```

<h1>Problem 3: Counting data</h1>
<b>This program tests your understanding of string manipulation and python dictionaries</b>

A website grabs biographical data from multiple web pages and collects relevant data in a python string. Each biographical entry is in a separate line, the lines are numbered (1,2,3,...), and data is separated by periods. Attributes are formatted in attribute_name: attribute_value pairs. An example of the string is below.

Write a function <span style="color:blue">summarize</span> that takes this text string and an attribute as an argument and returns a dictionary that contains the count of each value of the attribute. 

<h3>Example:</h3>
<pre>
sample_text = """1. John Smith. School: Harvard. Year: 2010. Major: History.
2. Name: Jill Smythe. School: Columbia. Year: 2010. Major: Psychology.
3. Name: Qing Sheng. School: Columbia. Year: 2010. Major: History.
4. Name: Abhishek Malhotra. School: Columbia. Year: 2015. Major: History.
5. Name: Pablo Balsera. School: Yale. Year: 2012. Major: Analytics.
6. Name: Overbrook Bha. School: Harvard. Year: 2011. Major: Analytics.
7. Name: Lidiya Guriyeva. School: Columbia. Year: 2010. Major: History.
8. Name: Cynthia Murphy. School: Columbia. Year: 2010. Major: History."""
</pre>

the function call:

<pre>
summarize(sample_text,"School")
</pre>

should return:
<pre>
{'Harvard': 2, 'Columbia': 5, 'Yale': 1}
</pre>

and 
<span style="color:blue">summarize(sample_text,"Year")</span>

should return:
<pre>
{'2010': 5, '2015': 1, '2012': 1, '2011': 1}
</pre>

There are multiple ways to do this but here's a list of useful python functions:
<li><a href="https://python-reference.readthedocs.io/en/latest/docs/str/split.html">split a string</a> use this to create a list of bio entries ("\n" is the end of line character)</li>
<li><a href="https://www.w3schools.com/python/ref_string_find.asp">find</a> finds the location of the first instance of a substring inside a string. Use this to find the starting location of the attribute and the location of the period following the end of the value of the attribute</li>
<li><a href="https://pymotw.com/2/collections/counter.html">Counter</a> use this to get the count of each instance. Either this, or some combination of an empty dict and <a href="https://www.w3schools.com/python/ref_dictionary_get.asp">get</a>. Or whatever works for you!</li>
<li>make sure your function returns a dictionary object</li>

<!-- grade_id: prob3_solution, points: 0, grade: false, solution: true -->
```python
from collections import Counter

def summarize(sample_text,attribute):
    # YOUR CODE HERE
    to_count = []
    start_point = len(attribute) + 3
    text = sample_text.split('.')
    for char in text:
        if attribute in char:
            to_count.append(char[start_point:])
    return dict(Counter(to_count))
    # raise NotImplementedError()
```

<!-- grade_id: prob3_public_test1, points: 0.11, grade: true, solution: false -->
```python
#Test case 1
sample_text = """1. Name: John Smith. School: Harvard. Year: 2010. Major: History.
2. Name: Jill Smythe. School: Columbia. Year: 2010. Major: Psychology.
3. Name: Qing Sheng. School: Columbia. Year: 2010. Major: History.
4. Name: Abhishek Malhotra. School: Columbia. Year: 2015. Major: History.
5. Name: Pablo Balsera. School: Yale. Year: 2012. Major: Analytics.
6. Name: Overbrook Bha. School: Harvard. Year: 2011. Major: Analytics.
7. Name: Lidiya Guriyeva. School: Columbia. Year: 2010. Major: History.
8. Name: Cynthia Murphy. School: Columbia. Year: 2010. Major: History."""

assert summarize(sample_text,"School") == {"Harvard" : 2, "Columbia" : 5, "Yale" : 1}, summarize(sample_text,"School")
```

<!-- grade_id: prob3_public_test2, points: 0.11, grade: true, solution: false -->
```python
#Test case 2
assert summarize(sample_text,"Year") == {"2010" : 5, "2015" : 1, "2012": 1, "2011" : 1}, summarize(sample_text,"Year") 
```

<!-- grade_id: prob3_test1, points: 0.26, grade: true, solution: false -->
```python

```

<!-- grade_id: prob3_test2, points: 0.26, grade: true, solution: false -->
```python

```

<!-- grade_id: prob3_test3, points: 0.26, grade: true, solution: false -->
```python

```

<h1>Problem 4: Functions and function conditions</h1>
<p>
A web page presents summary information from data in a list (see example below for the format of the list). Write a function <span style="color:blue">summary_stats</span> that returns a summary statistic from this list. The function takes four arguments: the list, a function that calculates the summary statistic, the column number (0,..,n-1) on which to calculate the statistic, and a boolean function that can be applied to the list to select elements. 

<h3>Example</h3>
<pre>
sample_data = [("John Smith","M", 430000),
               ("Jill Smythe","F", 250000),
                ("Qing Sheng","F", 323122),
                ("Abhishek Malhotra","M", 782122),
                ("Pablo Balsera","M",  600233),
                ("Overbrook Bha","M", 134112),
                ("Lidiya Guriyeva","F", 430000),
                ("Cynthia Murphy","F", 430000),]
</pre>

the function call (return the mean of column 2 for cases where col 1 has the value "F"): 

<pre>
summary_stats(sample_data,statistics.mean,2,lambda x: x[1]=="F")
</pre>

should return:

<pre>
358280.5
</pre>

AND, the function call (return the lowest value in column 2)

<pre>
summary_stats(sample_data,min,2)
</pre>

should return:

<pre>
134112
</pre>

<h3>Notes:</h3>
<li><span style="color:blue">lambda functions</span> are anonymous (unnamed) python functions (see: <a href="https://www.w3schools.com/python/python_lambda.asp">https://www.w3schools.com/python/python_lambda.asp</a> or <a href="https://realpython.com/python-lambda/">https://realpython.com/python-lambda/</a>)</li>
<li><span style="color:blue">boolean functions</span> are functions that return True or False</li>
<li><a href="https://www.w3schools.com/python/ref_func_filter.asp">filter</a> is a python standard library function that takes two arguments: a boolean condition (or a boolean lambda function) and a list. Filter applies the condition to each element in the list and returns only those elements that satisfy the condition. Use filter to apply the condition passed to the summary stats function, and then apply the statistic function to the result. Return the result</li>
<li>Your code MUST use a default value for the boolean condition function that selects all elements in the list (otherwise, the lowest value example will not work!)</li>

    




<!-- grade_id: prob4_solution, points: 0, grade: false, solution: true -->
```python
def summary_stats(sample_data,statistic,col,cond=lambda x: x):
    # YOUR CODE HERE
    filted = filter(cond, sample_data)
    to_operate = []
    for data in filted:
        to_operate.append(data[col])
    return statistic(to_operate)
    
    # raise NotImplementedError()
```

<!-- grade_id: prob4_public_test1, points: 0.25, grade: true, solution: false -->
```python
import statistics
sample_data = [("John Smith","M", 430000), 
("Jill Smythe","F", 250000), 
("Qing Sheng","F", 323122), 
("Abhishek Malhotra","M", 782122),
("Pablo Balsera","M",  600233), ("Overbrook Bha","M", 134112), 
("Lidiya Guriyeva","F", 430000), 
("Cynthia Murphy","F", 430000),]

# Test case 1
assert summary_stats(sample_data,statistics.mean,2,lambda x: x[1]=="F") == 358280.5, summary_stats(sample_data,statistics.mean,2,lambda x: x[1]=="F")
```

<!-- grade_id: prob4_public_test2, points: 0.25, grade: true, solution: false -->
```python
# Test case 2
assert summary_stats(sample_data,min,2) == 134112, summary_stats(sample_data,min,2)
```

```python

```